## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = "Tim Dentry"
# University of Arizona email address
EMAIL = "tdentry@email.arizona.edu"
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = "Andrew Quadro; we talked theory and ideas for approaches.  Jackson Mostoller; we traded ideas about how to construct feature dictionaries given the criteria"
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

# Mini Python tutorial

This course uses Python 3.8.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.8/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.8/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.8/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.8/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

Output:
```
h
e
l
l
o
[2, 3, 4, 5]
[2]
hello, Josuke!
Howdy, partner!
13
Hi, Fred!
[('radical', 4), ('analysis', 7), ('bighorn', 12), ('bounce', 32)]
[('analysis', 7), ('bighorn', 12), ('bounce', 32), ('radical', 4)]
```

```python

```

```python
from typing import Iterator, Sequence, Text, Tuple, Union

import numpy as np
from scipy.sparse import spmatrix, vstack
from sklearn.linear_model import LogisticRegression

import itertools
import pytest

# an NDArray is either a numpy array (ndarray) or a scipy sparse matrix (spmatrix)
NDArray  = Union[np.ndarray, spmatrix]
# type aliases for sequences of strings
# we'll use this type alias for our tokens
TokenSeq = Sequence[Text]
# ...and this one for our POS tags
TagSeq   = Sequence[Text]
```

```python
np.random.seed(42)
```

```python

```

<!-- grade_id: answer-imports, points: 0, grade: false, solution: true -->
```python
# Add your imports here (ex. classes from scikit-learn)
# YOUR CODE HERE
from sklearn.feature_extraction import DictVectorizer
from sklearn.preprocessing import LabelEncoder
from typing import Dict
from itertools import chain
from collections import Counter
```

## `read_ptbtagged`

First, we'll implement a function to read in our data.  Be careful!  Your first implementation might be off by 1...

<!-- grade_id: code-read-data, points: 0, grade: false, solution: true -->
```python
def read_ptbtagged(ptbtagged_path: str) -> Iterator[Tuple[TokenSeq, TagSeq]]:
    """
    Reads sentences from a Penn TreeBank .tagged file.
    Each sentence is a sequence of tokens and part-of-speech tags.

    Penn TreeBank .tagged files contain one token per line, with an empty line
    marking the end of each sentence. Each line is composed of a token, a tab
    character, and a part-of-speech tag. Here is an example:

        What	WP
        's	VBZ
        next	JJ
        ?	.

        Slides	NNS
        to	TO
        illustrate	VB
        Shostakovich	NNP
        quartets	NNS
        ?	.

    :param ptbtagged_path: The path of a Penn TreeBank .tagged file, formatted
    as above.
    :return: An iterator over sentences, where each sentence is a tuple of
    a sequence of tokens and a corresponding sequence of part-of-speech tags.
    """
    
    finalSequence=[]# Iterator
    
    with open(ptbtagged_path, "r") as file:
        lines = file.readlines()
        lineCount=len(lines) # counts the number of lines in the entire file being read
        tokenSequence = []
        posSequence = []
        for i, line in enumerate(lines): #counting the lines and parts
            parts = line.split('\t') #split on tab
            if(parts[0] in ["\n",""]):
                finalSequence.append ((tokenSequence, posSequence))
                tokenSequence = []
                posSequence = []
                # the above meets the condition of looping to the last line
                
            else:
                tokenSequence.append(parts[0]) # [What, 's, ...]
                posSequence.append(parts[1][:-1]) #[WP, VBZ,...]
                if (i == lineCount-1):  
                    finalSequence.append ((tokenSequence, posSequence))
                    # meets the condition of the last line being hit
          
    return finalSequence
```

```python
# identify part-of-speech tags from the Penn Treebank
PTB_TAGS = {
    "#", "$", "''", "``", ",", "-LRB-", "-RRB-", ".", ":", "CC", "CD", "DT",
    "EX", "FW", "IN", "JJ", "JJR", "JJS", "LS", "MD", "NN", "NNP", "NNPS",
    "NNS", "PDT", "POS", "PRP", "PRP$", "RB", "RBR", "RBS", "RP", "SYM", "TO",
    "UH", "VB", "VBD", "VBG", "VBN", "VBP", "VBZ", "WDT", "WP", "WP$", "WRB",
}
```

```python
# v = DictVectorizer(sparse=False)
# D = [{'token': 'Pierre', 'pos-1': '<s>'}, {'token': 'Vinken', 'pos-1': 'NNP'}, {'token': ',', 'pos-1': 'NNP'}, {'token': '61', 'pos-1': ','}, {'token': 'years', 'pos-1': 'CD'}, {'token': 'old', 'pos-1': 'NNS'}, {'token': ',', 'pos-1': 'JJ'}, {'token': 'will', 'pos-1': ','}, {'token': 'join', 'pos-1': 'MD'}, {'token': 'the', 'pos-1': 'VB'}, {'token': 'board', 'pos-1': 'DT'}, {'token': 'as', 'pos-1': 'NN'}, {'token': 'a', 'pos-1': 'IN'}, {'token': 'nonexecutive', 'pos-1': 'DT'}, {'token': 'director', 'pos-1': 'JJ'}, {'token': 'Nov.', 'pos-1': 'NN'}, {'token': '29', 'pos-1': 'NNP'}, {'token': '.', 'pos-1': 'CD'}]
# X = v.fit_transform(D)
# X

```

## Create some helper functions 

```python

```

## `Classifier`

Now we'll implement our classifier using the `Classifier` class.  Underlyingly, your classifier will make use of `sklearn` (see the attributes defined in the `__init__` method in the skeleton below).  

**HINT**: When encoding your features, be careful to only fit once!

<!-- grade_id: code-classifier, points: 0, grade: false, solution: true -->
```python
# Our MEMM
class Classifier:
    def __init__(self):
        """
        Initializes the classifier.
        """
        self.label_encoder = LabelEncoder()
        
        # Minimally, you must include the following features:
        # `token` (the current word) 
        # `pos-1` (the prior tag) 
        self.feature_encoder = DictVectorizer()
        self._feature_index = {}
        # multinomial logistic regression
        self.model = LogisticRegression(solver="liblinear", multi_class="ovr")
 


    def identify_features (self, tagged_sentences: Iterator[Tuple[TokenSeq, TagSeq]]):
        features = []
        for sen in tagged_sentences:  
            tokens = sen[0]
            pos = sen[1]
            tok_lis_len = len(tokens)
            
            for i in range(tok_lis_len):
                token = tokens[i]
                start = "<s>"
                if i > 0:
                    start = pos[i-1]
                features.append({"token":  token, "pos-1":  start})
        return features

    def train(self, tagged_sentences: Iterator[Tuple[TokenSeq, TagSeq]]) -> Tuple[NDArray, NDArray]:
        """
        Trains the classifier on the part-of-speech tagged sentences,
        and returns the feature matrix and label vector on which it was trained.

        The feature matrix should have one row per training token. The number
        of columns is up to the implementation, but there must at least be 1
        feature for each token, named "token=T", where "T" is the token string,
        and one feature for the part-of-speech tag of the preceding token,
        named "pos-1=P", where "P" is the part-of-speech tag string, or "<s>" if
        the token was the first in the sentence. For example, if the input is:

            What	WP
            's	VBZ
            next	JJ
            ?	.

        Then the first row in the feature matrix should have features for
        "token=What" and "pos-1=<s>", the second row in the feature matrix
        should have features for "token='s" and "pos-1=WP", etc. The alignment
        between these feature names and the integer columns of the feature
        matrix is given by the `feature_index` method below.

        The label vector should have one entry per training token, and each
        entry should be an integer. The alignment between part-of-speech tag
        strings and the integers in the label vector is given by the
        `label_index` method below.

        :param tagged_sentences: An iterator over sentences, where each sentence
        is a tuple of a sequence of tokens and a corresponding sequence of
        part-of-speech tags.
        
        :return: A tuple of (feature-matrix, label-vector).
        """
        X = list(tagged_sentences)
        
        pos_list = [pos for _, pos in X]
        label_vector = self.label_encoder.fit_transform(list(chain(*pos_list)))
    
        
        
        label_vector = self.label_encoder.fit_transform(list(chain(*pos_list)))
        feature_matrix = self.feature_encoder.fit_transform(self.identify_features(X))
        
        self.model.fit(feature_matrix, label_vector)
        
        self._feature_index = {feature_name: idx 
                             for idx, feature_name in 
                             enumerate(self.feature_encoder.feature_names_)}
        return feature_matrix, label_vector

        
                                                        
    def feature_index(self, feature: Text) -> int:
        """
        Returns the column index corresponding to the given named feature.

        The `train` method should always be called before this method is called.

        :param feature: The string name of a feature.
        
        :return: The column index of the feature in the feature matrix returned
        by the `train` method.
        """
        # may need to come back to this one
#         self.feature_index = {feat_name:  item for item, feat_name 
#                              in enumerate(self.feature_encoder.feature_names_)}
        return self._feature_index[feature]
                                        
        
    def label_index(self, label: Text) -> int:
        """
        Returns the integer corresponding to the given part-of-speech tag

        The `train` method should always be called before this method is called.

        :param label: The part-of-speech tag string.
        
        :return: The integer for the part-of-speech tag, to be used in the label
        vector returned by the `train` method.
        """
        if label =="<s>":
            return len(self.label_encoder.classes_)
        return self.label_encoder.transform([label])[0]

    def predict(self, tokens: TokenSeq) -> TagSeq:
        """
        Predicts part-of-speech tags for the sequence of tokens.

        This method delegates to either `predict_greedy` or `predict_viterbi`.
        The implementer may decide which one to delegate to.

        :param tokens: A sequence of tokens representing a sentence.
        
        :return: A sequence of part-of-speech tags, one for each token.
        """
        _, pos_tags = self.predict_greedy(tokens)
        # _, _, pos_tags = self.predict_viterbi(tokens)
        return pos_tags
    
    def more_features(self, tokens):
        f_matrix = []
        pos = []
        start_pos = "<s>"
        pos_t_2 = "<s>"
        tok_1 = "<starting_state>"
        
        for i, token in enumerate(tokens):
            if i > 0:
                tok_1 = tokens[i - 1]
                
            vect = self.feature_encoder.transform([{"token": token, "token-1": tok_1,
                "pos-1": start_pos,"pos-2": pos_t_2}])
            
            if i > 1:
                pos_t_2 = start_pos
            
            tag_prediction = self.model.predict(vect)
            start_pos = self.label_encoder.inverse_transform(tag_prediction)[0]
            
            f_matrix.append(vect.todense())
            pos.append(start_pos)

        return np.vstack(f_matrix), pos
    


    
        
    def predict_greedy(self, tokens: TokenSeq) -> Tuple[NDArray, TagSeq]:
        """
        Predicts part-of-speech tags for the sequence of tokens using a
        greedy algorithm, and returns the feature matrix and predicted tags.

        Each part-of-speech tag is predicted one at a time, and each prediction
        is considered a hard decision, that is, when predicting the
        part-of-speech tag for token i, the model will assume that its
        prediction for token i-1 is correct and unchangeable.

        The feature matrix should have one row per input token, and be formatted
        in the same way as the feature matrix in `train`.

        :param tokens: A sequence of tokens representing a sentence.
        
        :return: The feature matrix and the sequence of predicted part-of-speech
        tags (one for each input token).
        """
        return self.more_features(tokens)


        
        


    


    
      
    
    
    
    
    

```

```python
# part-of-speech tags from the Penn Treebank
PTB_TAGS = {
    "#", "$", "''", "``", ",", "-LRB-", "-RRB-", ".", ":", "CC", "CD", "DT",
    "EX", "FW", "IN", "JJ", "JJR", "JJS", "LS", "MD", "NN", "NNP", "NNPS",
    "NNS", "PDT", "POS", "PRP", "PRP$", "RB", "RBR", "RBS", "RP", "SYM", "TO",
    "UH", "VB", "VBD", "VBG", "VBN", "VBP", "VBZ", "WDT", "WP", "WP$", "WRB",
}
```

```python
test_dict = {"a": 1, "b": 2, "c":3}
test_dict.keys()
```

Output:
```
dict_keys(['a', 'b', 'c'])
```

## Test `.read_ptbtagged()` (3 pts)

Tests that you read in a) the correct number of sentences and tokens from the training data, b) that all `PTB_TAGS` were found in that partition of the data, and c) each token has exactly one corresponding tag.

<!-- grade_id: test-read-data, points: 3, grade: true, solution: false -->
```python
def test_read_ptbtagged():
    # keep a counter here (instead of enumerate) in case the iterator is empty
    token_count = 0
    sentence_count = 0
    for sentence in read_ptbtagged("data/PTBSmall/train.tagged"):
        assert len(sentence) == 2
        tokens, pos_tags = sentence
        assert len(tokens) == len(pos_tags)
        assert all(pos in PTB_TAGS for pos in pos_tags)
        token_count += len(tokens)
        sentence_count += 1
    assert token_count == 191969
    assert sentence_count == 8020

    # check the sentence count in the dev set too
    found = sum(1 for _ in read_ptbtagged("data/PTBSmall/dev.tagged")) 
    expected = 3106
    assert found == expected, f"Expected {expected} sentences, but found {found}"
    
test_read_ptbtagged()
```

<!-- grade_id: no-peeking, points: 5, grade: true, solution: false -->
```python
def no_peeking() -> bool:
    train_data = set(tuple(s) for s, tags in read_ptbtagged("data/PTBSmall/train.tagged"))
    dev_data   = set(tuple(s) for s, tags in read_ptbtagged("data/PTBSmall/dev.tagged"))
    if len(train_data) < 10:
        print("Something is wrong with the training data")
        return False
    if len(dev_data) < 10:
        print("Something is wrong with the dev data")
        return False
    
    for dev_ex in dev_data:
        if dev_ex in train_data:
            print(dev_ex)
            print("Dev data should not overlap with train data!")
            return False
    return True

# ensure the train and dev data is well-formed
assert no_peeking() is True, "Problem with train and/or dev data"
```

## Test features (5 pts)

This test ensures you are, per the definition of MEMM, minimally representing **token** and **prior tag** ($t_{i-1}$) features.  

Use the special symbol `<s>` to represent the prior tag of the first token in a sequence.

<!-- grade_id: test-features, points: 5, grade: true, solution: false -->
```python
def test_feature_vectors():
    clf       = Classifier()
    ptb_train = read_ptbtagged("data/PTBSmall/train.tagged")
    ptb_train = itertools.islice(ptb_train, 2)  # just the first 2 sentences
    features_matrix, labels_vector = clf.train(ptb_train)
    # num. tokens
    assert features_matrix.shape[0] == 31
    assert labels_vector.shape[0] == 31

    # train.tagged starts with
    # Pierre	NNP
    # Vinken	NNP
    # ,	,
    # 61	CD
    # years	NNS
    # old	JJ
    assert features_matrix[4, clf.feature_index("token=years")] == 1
    assert features_matrix[4, clf.feature_index("token=old")] == 0
    assert features_matrix[4, clf.feature_index("pos-1=CD")] == 1
    assert features_matrix[4, clf.feature_index("pos-1=NNS")] == 0
    assert features_matrix[0, clf.feature_index("pos-1=<s>")] == 1
    assert labels_vector[3] == clf.label_index("CD")
    assert labels_vector[4] == clf.label_index("NNS")
test_feature_vectors()
```

## Test greedy decoding (5pts)

In the greedy decoding approach, each tag is predicted one at a time, and each prediction is considered a **hard** decision.  In other words, when predicting the tag for token $t_{i}$, the model will assume that its prediction for the prior token $t_{i-1}$ is correct and unchangeable.

<!-- grade_id: test-greedy-decoding, points: 5, grade: true, solution: false -->
```python
def test_predict_greedy():
    clf        = Classifier()
    ptb_train  = read_ptbtagged("data/PTBSmall/train.tagged")
    ptb_train  = itertools.islice(ptb_train, 2)  # just the 1st 2 sentences
    clf.train(ptb_train)

    tokens = "Vinken is a director .".split()
    features_matrix, pos_tags = clf.predict_greedy(tokens)

    # check that there is one feature vector per POS tag
    assert features_matrix.shape[0] == len(pos_tags)

    # check that all POS tags are in the PTB tagset
    assert all(pos_tag in PTB_TAGS for pos_tag in pos_tags)

    def last_pos_index(ptb_tag):
        return clf.feature_index("pos-1=" + ptb_tag)

    # check that the first word ("The") has no pos-1 feature
    for ptb_tag in {"NNP", ",", "CD", "NNS", "JJ", "MD", "VB", "DT", "NN", "IN",
                    "VBZ", "VBG"}:
        assert features_matrix[0, last_pos_index(ptb_tag)] == 0

    # check that the remaining words have the correct pos-1 features
    for i, pos_tag in enumerate(pos_tags[:-1]):
        assert features_matrix[i + 1, last_pos_index(pos_tag)] > 0

test_predict_greedy()
```

## Minimum accuracy (4pts)

Your model should achieve >= 92% acccuracy against the first 100 sentences of the Penn Treebank development partition.  To achieve this accuracy, you may need to include additional contextual features (i.e., features that represent information about the surrounding words and/or tags).

**WARNING**: _this test may be slow to run (2 min.+)_

<!-- grade_id: test-min-accuracy, points: 4, grade: true, solution: false -->
```python
def test_accuracy():
    clf       = Classifier()
    ptb_train = read_ptbtagged("data/PTBSmall/train.tagged")
    clf.train(ptb_train)

    total_count   = 0
    correct_count = 0
    ptb_dev = read_ptbtagged("data/PTBSmall/dev.tagged")
    first_n = 100
    ptb_dev = itertools.islice(ptb_dev, first_n)  # just the 1st n sentences
    for tokens, pos_tags in ptb_dev:
        total_count += len(tokens)
        predicted_tags = clf.predict(tokens)
        assert len(predicted_tags) == len(pos_tags)
        for predicted_tag, true_tag in zip(predicted_tags, pos_tags):
            if predicted_tag == true_tag:
                correct_count += 1
    accuracy = correct_count / total_count

    # print out performance
    sg = f"\n{accuracy:.1%} accuracy on first {first_n} sentences of PTB dev"
    print(sg)
    return accuracy

# ensure solution meets min. accuracy
min_accuracy = 0.92
accuracy = test_accuracy()
assert accuracy >= min_accuracy
```

Output:
```
89.0% accuracy on first 100 sentences of PTB dev
AssertionError                            Traceback (most recent call last)
/var/folders/pl/k80fpf9s4f9_3rp8hnpw5x0m0000gq/T/ipykernel_13924/3905444770.py in <module>
     26 min_accuracy = 0.92
     27 accuracy = test_accuracy()
---> 28 assert accuracy >= min_accuracy

AssertionError:
```

```python

```